# EEN060 - Applied object-oriented programming

Teacher: [Carlos Natalino](https://www.chalmers.se/en/staff/Pages/Carlos-Natalino-Da-Silva.aspx) / Examiner: [Paolo Monti](https://www.chalmers.se/en/staff/Pages/Paolo-Monti.aspx)

[Canvas course page](https://chalmers.instructure.com/courses/17328)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.

**Delete** the line `raise NotImplementedError()` from the problems that you solve.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

You may also have problems that will be manually graded (e.g., HTML/CSS/JS/pseudocode problems).
These problems **do not** have a respective *test cell*.

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.

**If you want, you can solve this programming assignment using Google Colab**

Link: https://colab.research.google.com/

Just copy the following line to a code line and run it.

---

Make sure you fill your name and collaborators (if any) below.

```python
NAME = "Yu-Han Hua"
COLLABORATORS = ""
```

---

# Assignment 3

In this assignment, students should practice:
 - how to use built-in functionalities
 - how to use standard Python modules
 - how to find and use external Python modules
 - how to define classes
 - how to instantiate and intialize objects
 - how to define attributes
 - how to define methods
 - how to use inheritance
 - how to use composition
 - how to use dataclasses
 - how to set HTML styles property

 **Observations:**
 For the problems in this assignment, you should:
 - use the command `import`, if needed, within the solution cell, every time you need to use some module.

 Reading material:
 - lecture notes
 - Online documentation:
   - Python online documentation: https://docs.python.org/3.7/tutorial/modules.html
   - Specific modules (for optional reading):
     - datetime: https://docs.python.org/3.7/library/datetime.html
     - math: https://docs.python.org/3.7/library/math.html
     - random: https://docs.python.org/3.7/library/random.html
     - statistics: https://docs.python.org/3.7/library/statistics.html

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
try:
    from utils import validate_python_code, validate_html
except:
    import sys
    print("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.",
          file=sys.stderr)
```

## Level 1

Level 1 problems are simple problems that required only the knowledge obtained in the content of the lecture and the lecture notes of the particular lecture.

### Problem 1.1

Create a Python function called `degree_cosine` that receives one parameter that represents the angle in degrees, and returns the cosine value.

You must use the module `math` to solve this problem.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
# %%writefile degree_cosine_solution.py
# import math as s


# def degree_cosine(angle: float) -> float:

#     rads = s.radians (angle)

#     cosine_value = s.cos(rads)
    
#     return(cosine_value)

# degree_cosine(45)



```

```python
%%writefile degree_cosine_solution.py
import math as s


def degree_cosine(angle: float) -> float:
    rads = s.radians(angle)
    cosine_value = s.cos(rads)
    return cosine_value
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from degree_cosine_solution import degree_cosine

degree = 100
cosine = degree_cosine(degree)
print(f'The cosine for {degree} degrees is {cosine}')
print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-tests-1-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import degree_cosine_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from degree_cosine_solution import degree_cosine
except:
    raise ValueError("The module does not have the necessary function!")
with open("degree_cosine_solution.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "import math" in content, "You must import the `math` module."
from nose.tools import assert_equal, assert_almost_equal
assert_equal(degree_cosine(0), 1)
assert_almost_equal(degree_cosine(180), -1.0)
assert_almost_equal(degree_cosine(232), -0.6156614753256582)
from utils import validate_python_code
validate_python_code("degree_cosine_solution.py")
print('tests passed', u'\u2713')
```

### Problem 1.2

Create a Python function called `mean_stdev_var` that receives a list of numbers as parameter.
The function should return the mean, standard deviation and the variance, computed using the module `statistics`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile mean_stdev_var_solution.py
import statistics as s

def mean_stdev_var_solution(list_of_numbers: list[str]) -> float:

    mean = s.mean(list_of_numbers)
    standard = s.stdev(list_of_numbers)
    variance = s.variance(list_of_numbers)
    
    return mean, standard, variance


```

```python
%autoreload 2
#uncomment the line(s) below to debug
from mean_stdev_var_solution import mean_stdev_var

numbers = [30.363535357282277, 90.42641156890944, 35.88893670584669, 42.35277419651142, 23.343145012254173, 21.336637050463093, 90.55195368974573, 39.636659242884996, 60.46048988550831, 62.13510520797335]
print(f'Numbers: {numbers}')
stats = mean_stdev_var(numbers)
print(f'Mean: {stats[0]}')
print(f'Standard deviation: {stats[1]}')
print(f'Variance: {stats[2]}')
print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import mean_stdev_var_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from mean_stdev_var_solution import mean_stdev_var
except:
    raise ValueError("The module does not have the necessary function!")
with open("mean_stdev_var_solution.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "import statistics" in content, "You must import the `statistics` module."
from nose.tools import assert_equal
assert len(mean_stdev_var([15.888513602891996, 48.076345585896455, 53.99263004262558, 4.3020754809705135, 93.3166374758885])) == 3, 'The return of the function must be a sequence with 3 elements.'
assert_equal(mean_stdev_var([15.888513602891996, 48.076345585896455, 53.99263004262558, 4.3020754809705135, 93.3166374758885]),
            (43.115240437654606, 35.038788726175284, 1227.7167153975483))
assert_equal(mean_stdev_var([2.5852904919683173, 69.74922108279277, 39.12274849382771, 35.397001409275205, 99.21459870342308]),
            (49.21377203625742, 36.704000234175474, 1347.1836331903532))
from unittest.mock import patch
with patch('statistics.stdev') as mock_stdev:
    mean_stdev_var([2.5852904919683173, 69.74922108279277])
mock_stdev.assert_called_once()
from utils import validate_python_code
validate_python_code("mean_stdev_var_solution.py")
print('tests passed', u'\u2713')
```

### Problem 1.3

Create a Python function called `random_time_generator` that randomly generates a time between 08:00:00 and 19:59:59.
The function should return a `time` object.
For the sake of simplicity, the function should consider only hour, minute and second.

You must use modules `random` and `datetime` in your answer.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# %%writefile random_time_generator_solution.py
# import datetime as dt
# import random as r


# def random_time_generator() -> dt.time:
#     hour: int = r.randint(8, 19)
#     minu: int = r.randint(0, 59)
#     sec: int = r.randint(0, 59)
#     time: dt.time = dt.time(hour, minu, sec)
#     return time
```

<!-- grade_id: cell-answer-1-3, points: 0, grade: false, solution: true -->
```python
#%%writefile random_time_generator_solution.py

import random as r
import datetime


def random_time_generator() -> datetime.time:  

    random_range = r.randrange((8*60*60), ((19*60*60)+(59*60)+(59)))

    time = datetime.datetime.fromtimestamp(random_range).time()
    
    return(time)

print(random_time_generator())
```

Output:
```
17:37:59
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from random_time_generator_solution import random_time_generator

print('Here are some random times:')
for i in range(5):
    print(random_time_generator())
print('Execution finished', u'\u2713')
```

Output:
```
Here are some random times:
15:53:38
13:48:16
19:32:00
18:13:23
10:53:47
Execution finished ✓
```

<!-- grade_id: cell-test-1-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import random_time_generator_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from random_time_generator_solution import random_time_generator
except:
    raise ValueError("The module does not have the necessary function!")
with open("random_time_generator_solution.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "import datetime" in content, "You must import the `datetime` module."
    assert "import random" in content, "You must import the `random` module."

import datetime
from nose.tools import assert_equal, assert_almost_equal
assert isinstance(random_time_generator(), datetime.time), 'The return of the function is not a datetime.time object.'
# testing the probability
import collections
count_hour = collections.defaultdict(int)
count_minute = collections.defaultdict(int)
count_second = collections.defaultdict(int)
number_samples = 100000
for i in range(100000):
    d = random_time_generator()
    count_hour[d.hour] += 1
    count_minute[d.minute] += 1
    count_second[d.second] += 1
for year, count in count_hour.items():
    assert_almost_equal(count / number_samples, 1/(19 - 7), places=2)
for month, count in count_minute.items():
    assert_almost_equal(count / number_samples, 1/59, places=2)
for day, count in count_second.items():
    assert_almost_equal(count / number_samples, 1/59, places=2)
from utils import validate_python_code
validate_python_code("random_time_generator_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.4

Create a Python class with the name `AreaCalculator`.
During initialization with the `.__init__()` method, the method should receive a number and assign the value passed to the attribute `radius`.

Then, create a method named `area()` that returns the area of the circle using the `radius` attribute.
If the `radius` attribute is `None` or $\leq 0$ the method `area()` should return `None`.

**Remember:** from the `math` module, import `pi`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# %%writefile area_calculator_class.py


# from math import pi
# from typing import Optional

# # om radius är none eller mindre eller lika med 0 returna noll
# class AreaCalculator:
#     def __init__(self, radius:int):
#         self.radius=radius

#     def area(self) -> Optional[float]:
#         if (self.radius is not None) and (self.radius > 0):
#             A: float = (self.radius**2)*pi
#             return(A)
#         else:
#             return None
```

<!-- grade_id: answer-1-5, points: 0, grade: false, solution: true -->
```python
%%writefile area_calculator_class.py
# solution cell

from math import pi
from typing import Optional
dddd

class AreaCalculator:nn
    
    
    def __init__(self, radius: int):
        self.radius = radius
    def area(self) -> Optional[float]:
        if (self.radius is not None) and (self.radius > 0):
            my_area: float = (self.radius ** 2) * pi
            return my_area
        return None

```

Output:
```
Overwriting area_calculator_class.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from area_calculator_class import AreaCalculator

radius = 15
c = AreaCalculator(radius)
print(f'The area of a circle with radius {radius} is {c.area()}')
print('Execution finished', u'\u2713')
```

Output:
```
The area of a circle with radius 15 is 706.8583470577034
Execution finished ✓
```

<!-- grade_id: tests-1-5, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import area_calculator_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from area_calculator_class import AreaCalculator
except:
    raise ValueError("The module does not have the necessary class!")
with open("area_calculator_class.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "from math import pi" in content, "You must import `pi` from the `math` module."
from nose.tools import assert_equal, assert_almost_equal
import math
assert 'AreaCalculator' in dir(), 'Class AreaCalculator does not exist!'
assert 'area' in dir(AreaCalculator), 'Method area() does not exist!'
c = AreaCalculator(None)
assert c.area() == None, 'Method area() should return None for invalid values.'
c = AreaCalculator(0)
assert c.area() == None, 'Method area() should return None for invalid values.'
c = AreaCalculator(-1)
assert c.area() == None, 'Method area() should return None for invalid values.'
c = AreaCalculator(1)
assert isinstance(c.area(), float) or isinstance(c.area(), int), 'The return type of the function area() does not match the problem description.'
assert_almost_equal(c.area(), math.pi)
c = AreaCalculator(6)
assert isinstance(c.area(), float) or isinstance(c.area(), int), 'The return type of the function area() does not match the problem description.'
assert_almost_equal(c.area(), 113.097, places=2)
from utils import validate_python_code
validate_python_code("area_calculator_class.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
c:\python\kursnamn\programming-assignments\assignment03.ipynb Cell 30' in <module>
     <a href='vscode-notebook-cell:/c%3A/python/kursnamn/programming-assignments/assignment03.ipynb#ch0000028?line=28'>29</a> assert_almost_equal(c.area(), 113.097, places=2)
     <a href='vscode-notebook-cell:/c%3A/python/kursnamn/programming-assignments/assignment03.ipynb#ch0000028?line=29'>30</a> from utils import validate_python_code
---> <a href='vscode-notebook-cell:/c%3A/python/kursnamn/programming-assignments/assignment03.ipynb#ch0000028?line=30'>31</a> validate_python_code("area_calculator_class.py")
     <a href='vscode-notebook-cell:/c%3A/python/kursnamn/programming-assignments/assignment03.ipynb#ch0000028?line=31'>32</a> print('tests passed', u'\u2713')

File c:\python\kursnamn\programming-assignments\utils.py:146, in validate_python_code(filename, **args)
    <a href='file:///c%3A/python/kursnamn/programming-assignments/utils.py?line=140'>141</a>     raise 
```

### Problem 1.5

Create a Python class named `Person`.
In the initialization method, the parameters are:
1. Name (mandatory)
2. role (mandatory)
3. person number (mandatory)
4. address (optional, `None` by default)
5. phone number (optional, `None` by default)

The initialization method should save these parameters as attributes.

The class should also have a `__repr__()` method that returns a string containing only the attributes that are not `None`, separated by `\t` characters.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# from typing import Optional

# class Person:
#     def __init__(
#         self, 
#         name:str, 
#         role:str, 
#         person_number:int, 
#         adress: Optional[str], 
#         phone: Optional[str]
#     ):
    
#         self.name=name
#         self.role=role
#         self.person_number=person_number
#         self.adress=adress
#         self.phone=phone

#     def __repr__(self) -> str:
#         _repr: str= f"{self.name},t\{self.role}t\{self.person_number}"
#         if self.adress is not None:
#             _repr += f"t\{self.adress}"
#         if self.adress is not None:
#             _repr +="t\{self.phone}"
#         return (_repr)
    
```

<!-- grade_id: answer-1-6, points: 0, grade: false, solution: true -->
```python
#%%writefile person_class.py
# implimentation
from typing import Optional


class Person:
    def __init__(
        self, 
        name: str, 
        role: str, 
        person_number: str, 
        address: Optional[str] = None, 
        phone: Optional[str] = None,
    ):
        self.name = name
        self.role = role
        self.person_number = person_number 
        self.address = address
        self.phone = phone

    def __rep__(self) -> str:
        _repr: str = f"{self.name}\t{self.role}\t{self.person_number}"
        if self.address is not None:
            _repr += f"\t{self.address}"
        if self.phone is not None:
            _repr += f"\t{self.phone}"
        return  _repr        
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from person_class import Person
person = Person('John Doe', 'Manager', '201414202014')
print(f'The Person object converts to string as {person}')
print('Execution finished', u'\u2713')
```

<!-- grade_id: tests-1-6, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import person_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from person_class import Person
except:
    raise ValueError("The module does not have the necessary class!")
from nose.tools import assert_equal
assert 'Person' in dir(), 'Class Person does not exist!'
p = Person('John Doe', 'Manager', '201414202014' )
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    s = str(p)
mock_print.assert_not_called()
assert_equal(str(p), 'John Doe\tManager\t201414202014')
p = Person('Jane Doe', 'Supervisor', '201212202012', 'Gibraltargatan')
assert_equal(str(p), 'Jane Doe\tSupervisor\t201212202012\tGibraltargatan')
p = Person('Jane Doe', 'Supervisor', '201212202012', 'Gibraltargatan', '123456')
assert_equal(str(p), 'Jane Doe\tSupervisor\t201212202012\tGibraltargatan\t123456')
p = Person('Jane Doe', 'Supervisor', '201212202012', '123456')
assert_equal(str(p), 'Jane Doe\tSupervisor\t201212202012\t123456')
from utils import validate_python_code
validate_python_code("person_class.py")
print('tests passed', u'\u2713')
```

### Problem 1.6

Create a Python class named `Student` that inherits from the class `Person` created in the previous problem.
The class `Student` should override the initialization method of the class `Person` and set the attribute `role` always with the value `"Student"`, removing the need to pass it as a parameter to the method.
The initialization method should also include the parameter `program` as a mandatory one, before the address.

Finally, the class `Student` should also have a method `.print()` that uses the return of the method `.__repr__()` of the *parent* class and complements it with the program, printing out the result.

*Since the class in problem this problem inherits from the class created in the previous problem, note that you need to import the `Person` class from the module created in the previous problem in the solution cell of this one.*

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: answer-1-7, points: 0, grade: false, solution: true -->
```python
%%writefile student_class.py
# solution cell

from person import person

class Student:
    def __init__(self, role):
    self.role = role



```

```python
%autoreload 2
# uncomment the line(s) below to debug
from student_class import Student

s = Student('John Doe', '201414202014', 'TIEPL')
print('Student:')
s.print()
print('Execution finished', u'\u2713')
```

<!-- grade_id: tests-1-7, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import student_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from student_class import Student
except:
    raise ValueError("The module does not have the necessary class!")
from nose.tools import assert_equal
assert 'Student' in dir(), 'Class Student does not exist!'
assert 'print' in dir(Student), 'Class Student does not have a method `print`.'
s = Student('John Doe', '201414202014', 'TIEPL')
assert isinstance(s, Person), 'The class Student is not a child class of the class Person'
from unittest.mock import patch
with patch('student_class.print') as mock_print:
    s.print()
mock_print.assert_called_once_with('John Doe\tStudent\t201414202014\tTIEPL')
s = Student('John Doe', '201414202014', 'TIEPL', 'Gibraltargatan')
from unittest.mock import patch
with patch('student_class.print') as mock_print:
    s.print()
mock_print.assert_called_once_with('John Doe\tStudent\t201414202014\tGibraltargatan\tTIEPL')
assert hasattr(s, 'program'), 'The object does not seem to have the property `program`'
from utils import validate_python_code
validate_python_code("student_class.py")
print('tests passed', u'\u2713')
```

---
## Level 2

Problems in level 2 should build upon the content of the lecture and the content of previous lectures.
Therefore, the problems can be solved based solely on the lecture notes available for the course up to the lecture the assignment is about.

### Problem 2.1

Create a Python function called `lottery_numbers` that receives three parameters:
- how many numbers should be selected.
- the upper bound for the number generation;
- the lower bound for the number generation;

The function should generate a list of integer numbers with the correct number of elements, each element between the informed lower and upper bounds.
Remember that numbers in the list must not repeat.

The function should also validate the input. If any of the parameters is out of the range (i.e., lower bound should be $\geq 0$, upper bound should be $> 0$ and number should be $> 0$) the function should return `None`.


```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
%%writefile lottery_numbers_solution.py
# solution cell
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from lottery_numbers_solution import lottery_numbers
# 
# result = lottery_numbers(5, 50, 10)
# print(f'Generating 5 numbers between 10 and 50 is {result}')
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-2-1, points: 1, grade: true, solution: false -->
```python
# test cell
%autoreload 2
try:
    import lottery_numbers_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from lottery_numbers_solution import lottery_numbers
except:
    raise ValueError("The module does not have the necessary function!")
from nose.tools import assert_equal
result = lottery_numbers(5, 50, 10)
assert len(result) == len(set(result)), 'There are repeated elements in the list'
assert min(result) >= 10
assert max(result) <= 50
assert len(result) == 5
result = lottery_numbers(30, 100, 0)
assert len(result) == len(set(result)), 'There are repeated elements in the list'
assert min(result) >= 0
assert max(result) <= 100
assert len(result) == 30
assert_equal(lottery_numbers(30, 100, -1), None)
assert_equal(lottery_numbers(30, 0,0), None)
from utils import validate_python_code
validate_python_code("lottery_numbers_solution.py")
print('tests passed', u'\u2713')
```

### Problem 2.2

Create a Python class called `CourseRegister`.
In the initialization method, the class should receive the `name` and `code` of the course, and save them as attributes.

The class should also have a method called `.register()` that receives an object of the class `Student` and adds it to the list of enrolled students.
Note that this list should be initialized as an empty list in the initialization method, and called `list_registered_students`.
Note also that the objects should be added to the list only if they are an object of the `Student` class.
Finally, the method should not add the same student twice in the list.

The class should also have a method called `.remove()` that receives and object of the class `Student`, representing a student that is already enrolled, and removes it from the list of enrolled students.
Note that, before trying to remove the student from the list, the method should check whether the student is present in the list.

*Note that you should run your solutions to problems of the previous level for the test cell to work properly.*

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: answer-2-2, points: 0, grade: false, solution: true -->
```python
%%writefile course_register_class.py
# solution cell
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from course_register_class import CourseRegister
# c = CourseRegister('Applied object-oriented programming', 'EEN060')
# print(c.list_registered_students)
# s1 = Student('John Doe', '201414202014', 'TIEPL')
# c.register(s1)
# print(c.list_registered_students)
# p1 = Person('John Doe', '201414202014', 'Manager')
# c.register(p1)
# print(c.list_registered_students)
# s2 = Student('Jane Doe', '201212202012', 'TIEPL', 'Gibraltargatan')
# c.register(s2)
# print(c.list_registered_students)
# c.remove(s1)
# print(c.list_registered_students)
# print('Execution finished', u'\u2713')
```

<!-- grade_id: tests-2-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import course_register_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from course_register_class import CourseRegister
except:
    raise ValueError("The module does not have the necessary function!")
from nose.tools import assert_equal
assert 'CourseRegister' in dir(), 'Class CourseRegister does not exist!'
assert 'register' in dir(CourseRegister), 'Class CourseRegister does not have method `register`'
assert 'remove' in dir(CourseRegister), 'Class CourseRegister does not have method `remove`'
c = CourseRegister('Applied object-oriented programming', 'EEN060')
assert hasattr(c, 'name'), 'Class CourseRegister does not have property `name`.'
assert hasattr(c, 'code'), 'Class CourseRegister does not have property `code`.'
assert hasattr(c, 'list_registered_students'), 'Class CourseRegister does not have property `list_registered_students`.'
assert isinstance(c.list_registered_students, list), 'The attribute list_enrolled_students is not a list'
from student_class import Student
s1 = Student('John Doe', '201414202014', 'TIEPL')
c.register(s1)
assert_equal(len(c.list_registered_students), 1)
p1 = Person('John Doe', '201414202014', 'Manager')
c.register(p1)
assert_equal(len(c.list_registered_students), 1)
s2 = Student('Jane Doe', '201212202012', 'TIEPL', 'Gibraltargatan')
c.register(s2)
assert_equal(len(c.list_registered_students), 2)
c.remove(s1)
assert_equal(len(c.list_registered_students), 1)
c.remove(s1)
assert_equal(len(c.list_registered_students), 1)
c.remove(p1)
assert_equal(len(c.list_registered_students), 1)
c.register(s2)
assert_equal(len(c.list_registered_students), 1)
from utils import validate_python_code
validate_python_code("course_register_class.py")
print('tests passed', u'\u2713')
```

### Problem 2.3

Create a Python dataclass named `PersonAge`.
The attributes of this class are:
1. name (Type: sting)
2. job (Type: sting)
3. birthdate (Type: date)
4. age (Type: int, not an argument during initialization)

The attributes `name`, `job` and `birthdate` are mandatory during initialization.
Then, you need to create a method `.compute_age()` that uses the `birthdate` to compute how old the person is, saves it to the `age` property, and returns the `age` value.
Remember to import `dataclass` and `field` from `dataclasses` module.
Link to read more about data class: https://docs.python.org/3.9/library/dataclasses.html

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: solution-2-3, points: 0, grade: false, solution: true -->
```python
%%writefile person_age_class.py
# solution cell
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from datetime import date
# from person_age_class import PersonAge
# person1 = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
# person2 = PersonAge('Jane Doe', 'Manager', date(1980, 11, 15))
# print(f"The age of {person1.name} is {person1.compute_age()}")
# print(f"The age of {person2.name} is {person2.compute_age()}")
# print('Execution finished', u'\u2713')
```

<!-- grade_id: tests-2-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import person_age_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from person_age_class import PersonAge
except:
    raise ValueError("The module does not have the necessary class!")
import dataclasses
assert dataclasses.is_dataclass(PersonAge), "The class is not a dataclass!"
from datetime import date
from nose.tools import assert_equal
assert 'PersonAge' in dir(), 'Class PersonAge does not exist!'

p = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
assert hasattr(p, 'name'), 'Class PersonSorter does not have property `name`.'
assert hasattr(p, 'job'), 'Class PersonSorter does not have property `job`.'
p.compute_age()
assert hasattr(p, 'age'), 'Class PersonSorter does not have property `age`.'
assert p.age == 31, f"The computed age `{p.age}` is incorrect!"
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    s = str(p)
mock_print.assert_not_called()
person1 = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
person1.compute_age()
person2 = PersonAge('Jane Doe', 'Manager', date(1980, 11, 15))
person2.compute_age()
assert person1.age < person2.age, "The age computation seems incorrect!"
from utils import validate_python_code
validate_python_code("person_age_class.py")
print('tests passed', u'\u2713')
```

---
## Level 3: Pseudocode/HTML/CSS/JS problems

In these problems we are training the design part of the course, which will be used starting from week 6 in the course.

If the error messages are not very clear, you can use the online tool: https://onu2.s2.chalmers.se/nu/#textarea

Recommended sources:
- HTML: https://www.w3schools.com/html/
- CSS: https://www.w3schools.com/css/
- JavaScript: https://www.w3schools.com/js/

For this assignment, we study:
- lists: https://www.w3schools.com/html/html_lists.asp
- tables: https://www.w3schools.com/html/html_tables.asp

Visual Studio Code sometimes does not render HTML/CSS/JS as expected. We recommend testing your solutions here: https://www.w3schools.com/html/tryit.asp?filename=tryhtml_default



### Problem 3.1

Create a Python function called `list_generator` that receives a list of values.
The function should create a string that represents the HTML that shows the list of values, each one as an item of an unordered list.

The following example shows how the output of the test cell should look like.

Code for the test:
```
list_generator([
    'First item',
    'Second item',
    'Another item',
    'Final item'
])
```

Must result in the following result (the border is used just to delimit the output):
<html>
    <div style="border: 1px solid red;">
        <ul>
            <li>First item</li>
            <li>Second item</li>
            <li>Another item</li>
            <li>Final item</li>
        </ul>
    </div>
</html>

Note that each tag should be in one line, i.e., there should be a line break between every two tags.

<!-- grade_id: cell-html-answer, points: 0, grade: false, solution: true -->
```python
%%writefile list_generator_solution.py
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from list_generator_solution import list_generator
# print('HTML code:')
# print('You can copy the code below and paste in the W3Schools tryit editor to test it!')
# print(list_generator([
#     'First item',
#     'Second item',
#     'Another item',
#     'Final item'
# ]))
# print('Execution finished', u'\u2713')
```

```python
%autoreload 2
# DO NOT REMOVE THIS CODE >>> BEGINNING
# this cell assembles the HTML returned by the function that you wrote
# and renders it in the "preview" window
from IPython.core.display import display, HTML
from list_generator_solution import list_generator
html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body><p style="color: red;">start of your solution &gt;&gt;&gt;&gt;</p><hr/>
        """
html_header_offset = 1  # number of lines in the header
html_footer = """     <hr/><p style="color: red;">&lt;&lt;&lt;&lt; end of your solution</p></body>
</html>"""

html_to_test = html_header + list_generator([
    'First item',
    'Second item',
    'Another item',
    'Final item'
]) + html_footer
display(HTML(html_to_test))

# DO NOT REMOVE THIS CODE <<< END
```

<!-- grade_id: cell-html-test, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
from list_generator_solution import list_generator
import sys
from lorem_text import lorem
import requests
import random
from bs4 import BeautifulSoup
from utils import validate_html
has_error = False
html_header_offset = 1

html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body>
        """
html_footer = """
</html>"""

for i in range(5):

    num_items = random.randint(5, 20)

    items = []
    for i in range(num_items):
        items.append(lorem.words(random.randint(3, 20)))

    original_html = list_generator(items)

    html_to_test = html_header + original_html + html_footer

    soup = validate_html(html_to_test)

    final_html = ''
    for line in original_html.split('\n'):
        if line.strip() != '':
            final_html += line + '\n'

    html_to_test = html_header + final_html + html_footer

    soup = BeautifulSoup(html_to_test, 'html.parser')

    tag_set = set([str(tag.name) for tag in soup.find_all()])
    tag_pairs = []
    for tag_1 in tag_set:
        for tag_2 in tag_set:
            if tag_1 != tag_2:
                tag_pairs.append((tag_1, tag_2))

    for i, line in enumerate(final_html.split('\n')):
        for tag_1, tag_2 in tag_pairs:
            if f'<{tag_1}>' in line or f'</{tag_1}>' in line:
                assert f'<{tag_2}>' not in line and f'</{tag_2}>' not in line, f'Line {i+1} has two tags in the same line ({tag_1} and {tag_2}).\nYou must have them in different lines.\nLine {i+1}:\t{line}'

    tags_ul = soup.find_all('ul')
    assert len(tags_ul) == 1, f'The solution must have 3 unordered lists in your HTML. You have {len(tags_ul)} instead.'

    tags_li = tags_ul[0].find_all('li')
    assert len(tags_li) == num_items, f'The solution must have 14 list items in your HTML. You have {len(tags_li)} instead.'

    for idi, item in enumerate(tags_li):
        assert item.text == items[idi], f'The text in line {idi} is not the same as intended.'

from utils import validate_python_code
validate_python_code("list_generator_solution.py")

print('tests passed', u'\u2713')
```

### Problem 3.2

Create a Python function called `table_generator` that receives a dictionary.
The function should create a string that represents the HTML that shows the dictionary in a table with two columns: key and value.
Key and value should be explicitly defined as the heads of the table.

The following example shows how the output of the test cell should look like.

Code for testing:
```
table_generator(
    {'test key 1': 'test value 1',
    'test key 2': 'test value 2',
    'test key 3': 'test value 3',
    'test key 7': 'test value 7',
    'test key 6': 'test value 6',
    'test key 5': 'test value 5',
    'this_is_a_key': 'This is the value of the dictionary'})
```

Must result in the following output  (the border is used just to delimit the output):

<div style="border: 1px solid red;"><table><thead><tr><th>Key</th><th>Value</th></tr></thead><tbody><tr><td>testkey1</td><td>testvalue1</td></tr><tr><td>testkey2</td><td>testvalue2</td></tr><tr><td>testkey3</td><td>testvalue3</td></tr><tr><td>testkey7</td><td>testvalue7</td></tr><tr><td>testkey6</td><td>testvalue6</td></tr><tr><td>testkey5</td><td>testvalue5</td></tr><tr><td>this_is_a_key</td><td>Thisisthevalueofthedictionary</td></tr></tbody></table></div>


<!-- grade_id: answer-3-2, points: 0, grade: false, solution: true -->
```python
%%writefile table_generator_solution.py
# solution cell
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from table_generator_solution import table_generator
# print('HTML code:')
# s = table_generator(
#     {'test key 1': 'test value 1',
#     'test key 2': 'test value 2',
#     'test key 3': 'test value 3',
#     'test key 7': 'test value 7',
#     'test key 6': 'test value 6',
#     'test key 5': 'test value 5',
#     'this_is_a_key': 'This is the value of the dictionary'})
# print(s)
# print('Execution finished', u'\u2713')
```

```python
%autoreload 2
# DO NOT REMOVE THIS CODE >>> BEGINNING
# this cell assembles the HTML returned by the function that you wrote
# and renders it in the "preview" window

#from IPython import display
from table_generator_solution import table_generator
from IPython.core.display import display, HTML
html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body><p style="color: red;">start of your solution &gt;&gt;&gt;&gt;</p><hr/>
        """
html_footer = """     <hr/><p style="color: red;">&lt;&lt;&lt;&lt; end of your solution</p></body>
</html>"""
html = table_generator(
    {'test key 1': 'test value 1',
    'test key 2': 'test value 2',
    'test key 3': 'test value 3',
    'test key 7': 'test value 7',
    'test key 6': 'test value 6',
    'test key 5': 'test value 5',
    'this_is_a_key': 'This is the value of the dictionary'})
html_to_test = html_header + html + html_footer
display(HTML(html_to_test))

# DO NOT REMOVE THIS CODE <<< END
```

```python
%autoreload 2
# test cell
import sys
import requests
import random
from lorem_text import lorem
from bs4 import BeautifulSoup
from table_generator_solution import table_generator
has_error = False
html_header_offset = 1

html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body><p style="color: red;">start of your solution &gt;&gt;&gt;&gt;</p><hr/>
        """
html_footer = """     <hr/><p style="color: red;">&lt;&lt;&lt;&lt; end of your solution</p></body>
</html>"""

for i in range(5):

    input_dict = dict()
    num_keys = random.randint(5, 10)
    for i in range(num_keys):
        input_dict[lorem.words(random.randint(1, 3))] = lorem.words(random.randint(5, 10))

    original_html = table_generator(input_dict)

    html_to_test = html_header + original_html + html_footer
#     display(HTML(html_to_test))
    #print(html_to_test.count('\n'))
    assert html_to_test.count('\n') > 6, 'Your HTML seems to not have the sufficient lines.'

    soup = validate_html(html_to_test)

    final_html = ''
    for line in original_html.split('\n'):
        if line.strip() != '':
            final_html += line + '\n'
    html_to_test = html_header + final_html + html_footer

    soup = BeautifulSoup(html_to_test, 'html.parser')

    tag_set = set([str(tag.name) for tag in soup.find_all()])
    tag_pairs = []
    for tag_1 in tag_set:
        for tag_2 in tag_set:
            if tag_1 != tag_2:
                tag_pairs.append((tag_1, tag_2))

    for i, line in enumerate(final_html.split('\n')):
        for tag_1, tag_2 in tag_pairs:
            if f'<{tag_1}>' in line or f'</{tag_1}>' in line:
                assert f'<{tag_2}>' not in line and f'</{tag_2}>' not in line, f'Line {i+1} has two tags in the same line ({tag_1} and {tag_2}).\nYou must have them in different lines.\nLine {i+1}:\t{line}'

    tags_table = soup.find_all('table')
    assert len(tags_table) == 1, f'The HTML seems needs to have 1 table. Now it has {len(tags_table)} table(s).'

    tags_thead = tags_table[0].find_all('thead')
    assert len(tags_thead) == 1, 'The needs to have one `thead` tag.'

    tags_tr = tags_thead[0].find_all('tr')
    assert len(tags_tr) == 1, 'The table head needs one row.'

    tags_th = tags_tr[0].find_all('th')
    assert len(tags_th) == 2, f'The table must have 2 header columns. Now it has {len(tags_th)} column(s).'

    tags_tbody = tags_table[0].find_all('tbody')
    assert len(tags_tbody) == 1, 'The needs to have one `tags_tbody` tag.'

    tags_tr = tags_tbody[0].find_all('tr')
    assert len(tags_tr) == num_keys, f'The table must have {num_keys} rows. Now it has {len(tags_tr)} row(s).'

    for tag_tr in tags_tr:
        tags_td = tag_tr.find_all('td')
        assert len(tags_td) == 2, f'The table must have 2 columns. Now it has {len(tags_th)} column(s).\nRow with problems: \t{tag_tr}'

from utils import validate_python_code
validate_python_code("table_generator_solution.py")

print('tests passed', u'\u2713')
```